---
title        : Case study
description  : You know the in's and out's of two-table verbs with dplyr, but your knowledge is untried! Let's cement what you've learned with a real world application. 

--- type:VideoExercise lang:r xp: skills: key:0c3a3b2df4
## TBD 
TBD

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:2e43aab11e
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:095d42d564
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:23281fca6a
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:ea41628372
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:VideoExercise lang:r xp: skills: key:ea8cc2b0d8
## TBD 
TBD

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:27f2ed94c5
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:22792d1ad9
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:bdd590150b
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:e23d356fbd
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:VideoExercise lang:r xp: skills: key:9bc8ae8fbc
## TBD 
TBD

*** =video_link

--- type:NormalExercise lang:r xp: skills: key:2af6596b8e
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:8d94974a46
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:3dcf816f92
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang:r xp: skills: key:58dee97877
## TBD 
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct
