---
title        : Case study
description  : You know the in's and out's of two-table verbs with dplyr, but your knowledge is untried! Let's cement what you've learned with a real world application.

--- type:VideoExercise lang: xp: skills: key:a2f4cb7adf
## TBD
TBD

*** =video_link

--- type:NormalExercise lang: xp: skills: key:3c08f50873
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:94f67aa6a4
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:eab9dde2eb
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:a6ae8a9cf8
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:VideoExercise lang: xp: skills: key:79e29fc4ac
## TBD
TBD

*** =video_link

--- type:NormalExercise lang: xp: skills: key:a50f1c8386
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:a4884ef50b
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:0a56bd5603
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:9c216c4c80
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:VideoExercise lang: xp: skills: key:d134138e2a
## TBD
TBD

*** =video_link

--- type:NormalExercise lang: xp: skills: key:dd0b704ebf
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:f4f0f82bc9
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:3ac7e0dfa2
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct

--- type:NormalExercise lang: xp: skills: key:2dcba868e6
## TBD
TBD

*** =instructions

*** =hint

*** =pre_exercise_code

*** =sample_code

*** =solution

*** =sct
